Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Hai Dang Hoang"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    if len(minh)==0 and len(maxh)==0:
        maxh.append(elem*-1)
    else:
        if elem*-1 >= maxh[0]:
            heapq.heappush(maxh, elem*-1)
            if len(maxh)-1>len(minh):
                a = heapq.heappop(maxh)
                heapq.heappush(minh,a*-1)
        else:
            heapq.heappush(minh,elem)
            if len(minh) -1 > len(maxh):
                a = heapq.heappop(minh)
                heapq.heappush(maxh,a*-1)
    

```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    if len(minh) > len(maxh):
        median = minh[0]
    if len(maxh) > len(minh):
        median = maxh[0]*-1
    if len(minh)==len(maxh):
        median = (minh[0]+maxh[0]*-1)/2
    return median

```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2.0
3
4.0
5
6.0
7
8.0
9
10.0
11
12.0
13
14.0
15
16.0
17
18.0
19
20.0
21
22.0
23
24.0
25
26.0
27
28.0
29
30.0
31
32.0
33
34.0
35
36.0
37
38.0
39
40.0
41
42.0
43
44.0
45
46.0
47
48.0
49
50.0
```

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst case complexity happens when one heap's number of element is two elements larger than the other one after adding (which would trigger the popping and pushing process). Still, since both heappush and heappop have time complexity of O(lgn) due to the heap having lgn levels - and these two functions will go over at most all of those lgn levels, each level taking a constant number of steps. Therefore, the time complexity stays O(lgn) as they have the same dominant terms (with or without the extra process triggered).
To build a whole heap from scratch (adding elements to an initially blank list), the function will take O(lg1)+O(lg2)+O(lg2)+O(lg3)+ O(lg3)+...+O(lg(n/2))+ O(lg(n/2)), with a total number of elements of n (n/2 for each of max and min heap). The overall complexity of this is still O(lgn) as this is the dominant term.

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
The function median itself takes constant time. However, when building a whole heap, the whole process (which also includes add_to_median_heap) takes O(lgn)

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
This is better in terms of efficiency - the vanilla way of sorting the list would take nlogn time (heapsort), and taking the middle element will take constant time, to a total cost of O(nlogn). The above process only takes O(lgn) times, which is much faster.

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


```python
import random

eps = 1e-16 #just a very small number so that python does not count out of range
N = 10000
locations = [0.0, 0.5, 1.0 - eps] #the "tool" to find the location of the three numbers for the median


def median(x1, x2, x3): #find the middle number among three numbers - choosing the "median of 3"
    if (x1 < x2 < x3) or (x3 < x2 < x1):
        return x2
    elif (x1 < x3 < x2) or (x2 < x3 < x1):
        return x3
    else:
        return x1 #output the median among the three inputs

def qsort(lst,k): #the quick sort function. Takes a list as an input
    indices = [(0, len(lst))] #a list of tuples

    while indices: #while indices is not empty
        (frm, to) = indices.pop() #take out the last element in the list, assign them to two variables frm and to
        if frm == to: #if frm is equal to "to", meaning that the list is empty, end the while loop and continue after that
            continue

        # Find the partition:
        N = to - frm #number of elements in the array
        inds = [frm + int(N * n) for n in locations]   #find the index of the three numbers for the median
        values = [lst[ind] for ind in inds] #find the three numbers
        partition = median(*values) #find the median to partition

        # Split into lists:
        lower = [a for a in lst[frm:to] if a < partition] #the lower array, split with the newly found median
        upper = [a for a in lst[frm:to] if a > partition] #the upper array
        counts = sum([1 for a in lst[frm:to] if a == partition]) #this is to account for the fact that there may be duplicates of the median

        ind1 = frm + len(lower) #ind1 is the index of the end of the lower array
        ind2 = ind1 + counts #index 2 is the start of the upper array (plus counts to account for duplicates)

        # Push back into correct place:
        lst[frm:ind1] = lower #put back into the original list by substitution
        lst[ind1:ind2] = [partition] * counts
        lst[ind2:to] = upper

        # Enqueue other locations
        if k in range(ind1,ind2): #if the provided k is within the range of the index of the partition - meanning that we found our answer
            return lst[ind1]
            break
        elif k<ind1: #if the provided k is smaller than index 1, we will need to push the lower list up (since what we want is within this list)
            indices.append((frm, ind1)) #then append this back to the indices list (list of tuples above), to loop over
        elif k>=ind2:
            indices.append((ind2, to))
    return lst #return the sorted list
```

<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    if k > len(lst):
        print("This is not feasible")
    else:
        return qsort(lst,k)
    #raise NotImplementedError()
```

```python
qsort([1,2,3,4,5,6,7],0)
```

Output:
```
1
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))

```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
The above code does not use any recursion - instead, it uses a while loop (which technically recurs when the condition is not met, but is not strictly a recursion). The code resembles that of quick sort , however we only divide until we find the kth smallest element (so fewer number of levels), and the cost of each level is much less (since we will only push up the part of the list that contains the desired element). 
    
Assuming completely balanced splits, the cost of each level will be O(n), then O(n/2), O(n/4),... 

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
In the best case, k will simply be found in the first loop (which is either the last, middle or first element by using the above median of three method). In this case, quicksort runs once, with cost of O(n)

## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
In the worst case, k can only be found in the last loop - meaning that the full number of levels of lg(n) will be run. The total cost will be O(n) + O(n/2) + O(n/4) + ... + O(1) = O(n) - since the dominant term is still n.

```python

```